---
title       : Esimene kodutöö
description : Esimese nädala ülesanded
--- type:NormalExercise lang:r xp:100 skills:1 key:07547c77c0
## Arvutamine


*** =instructions
Lahenda ülesanded

1. Leia arvu 25 kümnendlogaritm.
2. Leia arvu 25 naturaallogaritm.
3. Liida arvud ühest kuni 250-ni: $\sum_{i=1}^{250} i$.

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}
# Leia arvu 25 kümnendlogaritm


# Leia arvu 25 naturaallogaritm


# Liida arvud ühest kuni 250-ni 



```

*** =solution
```{r}
# Leia arvu 25 kümnendlogaritm
log10(25)

# Leia arvu 25 naturaallogaritm
log(25)

# Liida arvud ühest kuni 250-ni 
sum(1:25)

```

*** =sct
```{r}

```
